Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Lars Janssen"
```

---

For those not familiar with Python, a quick overview is given [here](https://github.com/palcu/python-for-competitive-programming/blob/master/python-for-competitive-programming.ipynb).

# Notebook BAPC week 1: Input and output

## These notebooks
In this course, we will work with Jupyter Notebooks. These notebooks allow us to mix regular text with code, and are therefore very useful for homework and exercises. You will receive one notebook weekly, and your notebook will be graded automatically using so-called *test cells*. These cells are marked with `# TEST`, and are *read-only*, which means that any changes you make will be overwritten during grading. There will be tests that are visible to you, and also *hidden* tests that may test edge cases in your code. You will receive points only when all visible and hidden tests pass.

You will receive 1 point automatically when the following cell passes.

<!-- grade_id: cell-a4e5c0ffeb96dab9, points: 1, grade: true, solution: false -->
```python
# TEST that our life is not a lie
assert True
assert not False
```

## Notebook setup

First, we want to make sure that we are actually running Python 3. Run the following cell by pressing Shift-Enter. If you get an `AssertionError`, please raise your hand.

```python
# Run the following piece of code!

# This asserts that we are running Python 3.
import sys
assert sys.version_info >= (3,)
```

```python
# Run the following piece of code!

# This allows us to emulate command line I/O.
# Don't worry about how it works.
from contextlib import redirect_stdout
from io import StringIO
from sys import stdin
# Overwrite the jupyter input function.
def input():
    return stdin.readline()
```

## Exercise 1: some basic Python 3

In this course, we will be building Python programs that take *input* from *standard in* and produce *output* to *standard out*. Reading a line from stdin (standard in) can be done using the `input()` function, and output is formed using the `print` function:

```python
# This emulates command-line input.
stdin = StringIO("Some line of input")

line = input()
print("We read the following input:", line)
```

Output:
```
We read the following input: Some line of input
```

When running a stand-alone Python script in a terminal, the command `input()` reads one line of input that you type in the terminal. The command `print()` displays text in this terminal.

If you have an input file `my_input.in`, you can run `python3 my_script.py < my_input.in`, and it will be like the text from `my_input.in` is typed into the terminal. Then the command `input()` will read one line of `my_input.in` every time it is called.

### Basic string manipulation

```python
# Given a `line` of input, we can convert all letters to lowercase:
line_lowercase = line.lower()
print(line)
print(line_lowercase)

# Or split the line at every space:
words = line_lowercase.split(" ")
print(words)
```

Output:
```
Some line of input
some line of input
['some', 'line', 'of', 'input']
```

```python
# The string type has a *lot* of useful methods,
# which we can inspect using the `help` function:
help(type(line))
```

Output:
```
Help on class str in module builtins:

class str(object)
 |  str(object='') -> str
 |  str(bytes_or_buffer[, encoding[, errors]]) -> str
 |  
 |  Create a new string object from the given object. If encoding or
 |  errors is specified, then the object must expose a data buffer
 |  that will be decoded using the given encoding and error handler.
 |  Otherwise, returns the result of object.__str__() (if defined)
 |  or repr(object).
 |  encoding defaults to sys.getdefaultencoding().
 |  errors defaults to 'strict'.
 |  
 |  Methods defined here:
 |  
 |  __add__(self, value, /)
 |      Return self+value.
 |  
 |  __contains__(self, key, /)
 |      Return key in self.
 |  
 |  __eq__(self, value, /)
 |      Return self==value.
 |  
 |  __format__(...)
 |      S.__format__(format_spec) -> str
 |      
 |      Return a formatted version of S as described by format_spec.
 |  
 |  __ge__(self, value, /)
 |      Return self>=value.
 |  
 |  __getattribute__(self, name, /)
 |      Return getatt
```

```python
# We can also index inside a string to receive a single character...
print("The first character is", line[0])

# ...and get its Unicode/ASCII code using `ord`:
print("The ASCII-code of", line[0], "is", ord(line[0]))
```

Output:
```
The first character is S
The ASCII-code of S is 83
```

```python
# We can use string formatting to produce formatted output.
# Here are three possibilities; choose whichever one you like.
print("The ASCII-code of %s is %d" % (line[0], ord(line[0])))
print("The ASCII-code of {} is {}".format(line[0], ord(line[0])))
print(f"The ASCII-code of {line[0]} is {ord(line[0])}") # Note the f right before the string!
```

Output:
```
The ASCII-code of S is 83
The ASCII-code of S is 83
The ASCII-code of S is 83
```

### Functions

```python
# Functions in Python are very readable and easy to define:
def some_function(argument1, argument2):
    if argument1 > argument2:
        return 'first is largest'
    else:
        return 'second is largest'

# You can usually apply a function to all types of arguments that make sense.
print(some_function(100.0, 101.0))
print(some_function(15, 12))
print(some_function("def", "abc"))
```

Output:
```
second is largest
first is largest
first is largest
```

```python
# We can use assertions to test our function:
assert some_function(100, 99.0) == 'first is largest'
assert some_function(100, 101.0) == 'second is largest'
```

In the following cell, finish the function `letter_of_alphabet`. The test-cell below it is used for grading.

<!-- grade_id: cell-d7e91f5fc0c5150b, points: 0, grade: false, solution: true -->
```python
def letter_of_alphabet(character):
    """ Returns the number of this letter in the alphabet. """
    assert type(character) == str and len(character) == 1
    assert 'a' <= character <= 'z' or 'A' <= character <= 'Z'
    if('a' <= character <= 'z'):
        return (ord(character) - 96)
    if('A' <= character <= 'Z'):
        return (ord(character) - 64)
    raise NotImplementedError()
```

<!-- grade_id: cell-c4878bbb96fafbe3, points: 1, grade: true, solution: false -->
```python
# TEST that the function works as intended.
assert letter_of_alphabet('a') == 1
assert letter_of_alphabet('z') == 26
assert letter_of_alphabet('A') == 1, "Does your solution also accept uppercase numbers?"
```

### Numbers

In competitive coding, input is always given to us in string format. However, the vast majority of the time, we need either integers or floats. We can freely convert between these types using the `str`, `int`, and `float` functions:

```python
word_str = "3"
word_int = int(word_str)
word_flt = float(word_str)
print(word_str, type(word_str))
print(word_int, type(word_int))
print(word_flt, type(word_flt))
```

Output:
```
3 <class 'str'>
3 <class 'int'>
3.0 <class 'float'>
```

```python
# In Python, the division operator `/` produces a float:
print(3/2)

# The double-slash is C-style integer-division:
print(3//2)

# We can print integers or floats using C-style formatting:
print("%d is equal to %d but not equal to %g" % (3/2, 3//2, 3/2))
```

Output:
```
1.5
1
1 is equal to 1 but not equal to 1.5
```

```python
# Computing powers is also very easy, much easier than in C:
print(2 ** 10)
print(2 ** 100)
print(len(str(2 ** 10000)))
```

Output:
```
1024
1267650600228229401496703205376
3011
```

### Lists

```python
some_list = [4,2,5,3]
print(some_list)

# We can index with negative numbers to receive indices from the right:
print(some_list[0] * some_list[-1])

# We can append to a list like so:
some_list.append(0)
print(some_list)

# We can sort the list like so:
sorted_list = sorted(some_list)
print(sorted_list)
```

Output:
```
[4, 2, 5, 3]
12
[4, 2, 5, 3, 0]
[0, 2, 3, 4, 5]
```

```python
# We can loop over things in a list like so:
for x in some_list:
    print(x, end=' ')  # Print no newline at the end.
print()  # Print a quick newline.

# Or loop over the "reversed" list like so:
for x in reversed(some_list):
    print(x, end=' ')
print()  # Print a quick newline.

# A reversed list represents a list but is technically not one...
print(reversed(some_list), type(reversed(some_list)) == list)

# So we can convert it to a list explicitly:
reversed_list = list(reversed(some_list))
print(reversed_list, type(reversed_list) == list)
```

Output:
```
4 2 5 3 0 
0 3 5 2 4 
<list_reverseiterator object at 0x7fb97c2a1908> False
[0, 3, 5, 2, 4] True
```

```python
# There are other special list-like things, such as `range`:
range_10 = range(10)

# You mostly use `range` for making standard for-loops.
# For example, we can use this `range` to initialize a list...
squares = []
for x in range(10):
    squares.append(x ** 2)
print(squares)
```

Output:
```
[0, 1, 4, 9, 16, 25, 36, 49, 64, 81]
```

```python
# We could also use a *list comprehension*:
squares = [x ** 2 for x in range(10)]
print(squares)

# we can even use a list comprehension to do more complicated stuff:
print([x ** 2 for x in range(10) if x % 2 == 1])

# or just to initialize a list with zeros:
zeros = [0 for x in range(5)]
print(zeros)
```

Output:
```
[0, 1, 4, 9, 16, 25, 36, 49, 64, 81]
[1, 9, 25, 49, 81]
[0, 0, 0, 0, 0]
```

One usecase of list comprehensions is to load a string into a list of integers. Complete the function below.

<!-- grade_id: cell-deb47a38c51b4a12, points: 0, grade: false, solution: true -->
```python
def to_list_of_integers(string):
    strinput = string.split(" ")
    intinput = [int(strinput[i]) for i in range(len(strinput))]
    return intinput
    raise NotImplementedError()
```

<!-- grade_id: cell-eebeac3aac82d0e3, points: 1, grade: true, solution: false -->
```python
# TEST that the function does what it should.
assert to_list_of_integers('0 101010') == [0, 101010]
assert to_list_of_integers('6 4 5 3 5') == [6, 4, 5, 3, 5]
```

### Modules
Python contains a *lot* of very useful modules from which standard functions can be imported. For now, let's consider the `math` module. Scroll through the list of functions (and constants!) it provides.

There are modules for almost anything you can think of, but in programming contests you are only allowed to use the modules that come with "vanilla Python". So usually no `numpy`, for example.

```python
import math
help(math)
```

Output:
```
Help on built-in module math:

NAME
    math

DESCRIPTION
    This module is always available.  It provides access to the
    mathematical functions defined by the C standard.

FUNCTIONS
    acos(...)
        acos(x)
        
        Return the arc cosine (measured in radians) of x.
    
    acosh(...)
        acosh(x)
        
        Return the inverse hyperbolic cosine of x.
    
    asin(...)
        asin(x)
        
        Return the arc sine (measured in radians) of x.
    
    asinh(...)
        asinh(x)
        
        Return the inverse hyperbolic sine of x.
    
    atan(...)
        atan(x)
        
        Return the arc tangent (measured in radians) of x.
    
    atan2(...)
        atan2(y, x)
        
        Return the arc tangent (measured in radians) of y/x.
        Unlike atan(y/x), the signs of both x and y are considered.
    
    atanh(...)
        atanh(x)
        
        Return the inverse hyperbolic tangent of x.
    
    ceil(...)
        ceil(x)
        
 
```

```python
print(math.pi, math.e, math.sqrt(3) == 3 ** 0.5)
```

Output:
```
3.141592653589793 2.718281828459045 True
```

Consider the following [xkcd](https://xkcd.com/1047/), and in particular the bottom-most bit.

Use the cell below to compute, for the relations $L = R$ below, the error $|L - R|$:
* $\sqrt{2} = \tfrac{3}{5} + \tfrac{\pi}{7 - \pi}$,
* $\cos(\tfrac{\pi}{7}) + \cos(\tfrac{3 \pi}{7}) + \cos(\tfrac{5 \pi}{7}) = \tfrac{1}{2}$, and 
* $\sqrt{5} = \tfrac{13 + 4 \pi}{24 - 4 \pi}$.

<!-- grade_id: cell-24b64fccdb789804, points: 0, grade: false, solution: true -->
```python
def xkcd_approximations():
    first = math.fabs(math.sqrt(2) - (3/5 + (math.pi)/(7-math.pi)))
    second = math.fabs(math.cos(math.pi/7) + math.cos((3*math.pi)/7) + math.cos((5*math.pi)/7) - 1/2)
    third = math.fabs(math.sqrt(5) - (13+4*math.pi)/(24-4*math.pi))
    return [first, second, third]
    raise NotImplementedError()
```

<!-- grade_id: cell-62b32c71aea86bd9, points: 1, grade: true, solution: false -->
```python
# TEST that we get the right approximations.
# NOTE: there will be hidden tests below to check your answer!
assert len(xkcd_approximations()) == 3
a, b, c = xkcd_approximations()

```

## Exercise 2: Kattis problem "Last Factorial Digit"
Let's begin with the trivial problem [Last Factorial Digit](https://open.kattis.com/problems/lastfactorialdigit). Read the problem statement. A solution to this problem consists of a (Python) program that receives some input, does some processing, and prints some output.

The function `lastfactorialdigit_solution` below solves a single testcase. It takes as arguments one integer ($N$) and returns an integer, the final digit of the factorial $N!$. Finish this function below.

<!-- grade_id: cell-460fd6512a79a9f1, points: 0, grade: false, solution: true -->
```python
def lastfactorialdigit_solution(N):
    if(N < 4):
        return (math.factorial(N))
    elif(N == 4):
        return 4
    else:
        return 0
    raise NotImplementedError()
```

We can verify that the function returns the correct answer for the different test cases in `Sample Input 1`:

<!-- grade_id: cell-13e09f6310569942, points: 0, grade: true, solution: false -->
```python
# TEST that the solution gives the right answer to the first test case.
assert lastfactorialdigit_solution(1) == 1
assert lastfactorialdigit_solution(2) == 2
assert lastfactorialdigit_solution(3) == 6
```

In a Kattis scenario, we read input as text from stdin and we print to stdout. Use the function `lastfactorialdigit_solution` to complete the problem below.

<!-- grade_id: cell-6252388dba6b645f, points: 0, grade: false, solution: true -->
```python
def lastfactorialdigit():
    print(output[0])
    return output
    raise NotImplementedError()
```

<!-- grade_id: cell-ecb52cc01c6c8901, points: 1, grade: true, solution: false -->
```python
# TEST that the solution gives the expected result for the given sample inputs and outputs.
for (stdin, expected_out) in [(StringIO('3\n1\n2\n3\n'), '1\n2\n6\n'), (StringIO('2\n5\n2\n'), "0\n2\n")]:
    output = StringIO()
    with redirect_stdout(output):
        lastfactorialdigit()
    assert output.getvalue() == expected_out
```

Output:
```
<ipython-input-48-9ee91d22bc63> in <module>
      3     output = StringIO()
      4     with redirect_stdout(output):
----> 5         lastfactorialdigit()
      6     assert output.getvalue() == expected_out

<ipython-input-47-9a2577b3ae1f> in lastfactorialdigit()
      1 def lastfactorialdigit():
----> 2     print(output[0])
      3     return output
      4     raise NotImplementedError()

TypeError: '_io.StringIO' object is not subscriptable
```

```python

```